# Python Entry Test

## Part 3 of 4: Code Organization (Indentation, execution flow, import, functions)

## Instructions

Below are 5 Exercises: python problems you need to solve, in order to pass this test. For each correct answer, you will be assigned points, with a maximum of 100 points in total. If you score 80 points or higher, you have passed the test. You cannot see your score immediately; your test will be graded as part of a batch job, later on.

Note the following:
1. You are asked to code in **Python3** (version 3), not Python2
2. For some answers, you obtain a *skeleton*, to frame your answer. You are asked to complete it
3. The answer you are expected to given starts where it says:
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```
   Start your answer by *removing these two lines*. Do not put your answer anywhere else than *just here, in this box*
4. You may use as many lines for your answer as you see necessary
5. In general, however, no more than 5 lines of code are needed
6. Ignore all boxes that are completely empty; just focus on `# YOUR CODE HERE`

# Begin of the Test

#### Exercise 1: below is a piece of code that, if you run it, yields an `IndentationError`. You are asked to correct all the indentation of the source code, so that the program runs and produces the correct result: `code=10`. Do not remove, change or add any code; just repair the indentation.

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
    import math
    
 def square_root(x):
    return math.sqrt(x)

 def run():
      print('code={:d}'. \
            format(int(square_root(100))))
    
# call run() to run this code ...
run()
### BEGIN SOLUTION
### END SOLUTION
```

Output:
```
File "<ipython-input-3-492cd03a7fd3>", line 3
    def square_root(x):
    ^
IndentationError: unexpected indent
```

<!-- grade_id: at11, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):

    def test_import_indentation(self):
        ''' verify fimport math availlable '''
        try:
            okey = hasattr(math, 'pi' )
        except Exception:
            self.fail("test import and use math failed")
            
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_import_indentation (__main__.MyEntryTestClasses)
verify fimport math availlable ... FAIL

======================================================================
FAIL: test_import_indentation (__main__.MyEntryTestClasses)
verify fimport math availlable
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-4-67b86646ddd8>", line 9, in test_import_indentation
    okey = hasattr(math, 'pi' )
NameError: name 'math' is not defined

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "<ipython-input-4-67b86646ddd8>", line 11, in test_import_indentation
    self.fail("test import and use math failed")
AssertionError: test import and use math failed

----------------------------------------------------------------------
Ran 1 test in 0.003s

FAILED (failures=1)
```

<!-- grade_id: at12, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):

    def test_indentation_corrected_ok(self):
        ''' verify function square_root() is callable '''
        try:
            sqrt = square_root(100)
            assert sqrt == 10
        except Exception:
            self.fail("test indentation corrected Ok failed")
            
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS   
```

Output:
```
test_indentation_corrected_ok (__main__.MyEntryTestClasses)
verify function square_root() is callable ... FAIL

======================================================================
FAIL: test_indentation_corrected_ok (__main__.MyEntryTestClasses)
verify function square_root() is callable
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-5-2ebe7053843c>", line 9, in test_indentation_corrected_ok
    sqrt = square_root(100)
NameError: name 'square_root' is not defined

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "<ipython-input-5-2ebe7053843c>", line 12, in test_indentation_corrected_ok
    self.fail("test indentation corrected Ok failed")
AssertionError: test indentation corrected Ok failed

----------------------------------------------------------------------
Ran 1 test in 0.002s

FAILED (failures=1)
```

<!-- grade_id: at13, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):

    def test_correct_source_indentation(self):
        
        expected_string = 'code=10'   
        with captured_output() as (out, err):
            run()

        output = out.getvalue().strip()
        self.assertEqual(output, expected_string)
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS

```

Output:
```
test_correct_source_indentation (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_correct_source_indentation (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-6-6beeb4031bbf>", line 24, in test_correct_source_indentation
    run()
NameError: name 'run' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.002s

FAILED (errors=1)
```

#### Exercise 2: the below code is not printing anything. Why not? Repair it by adding exactly two lines of code.

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
def give_me_the_answer(text):
    print('{:s}'.format(text))
    
def prepare_the_answer(immediate):
    if immediate == True:
        answer = "here is your answer!"
    else:
        answer = ""
    return answer

# add a single line setting a variable ...
### BEGIN SOLUTION
NOW = True
### END SOLUTION

text = prepare_the_answer(NOW)

# add a single line here to produce the output ...
### BEGIN SOLUTION
give_me_the_answer(text)
### END SOLUTION
```

Output:
```
here is your answer!
```

<!-- grade_id: at21, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(NOW, True)
### END HIDDEN TESTS
```

<!-- grade_id: at22, points: 12, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):

    def test_printing_answer(self):
        
        expected_output = 'here is your answer!'
        
        with captured_output() as (out, err):
            give_me_the_answer(text)
    
        # This can go inside or outside the `with` block
        output = out.getvalue().strip()
        self.assertEqual(output, expected_output)
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_printing_answer (__main__.MyEntryTestClasses) ... ok

----------------------------------------------------------------------
Ran 1 test in 0.002s

OK
```

#### Exercise 3: in the below code skeleton, insert the necessary `import`'s as specified

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
# import math and name it by the alias m
# import numpy and name it by the alias np
# import nose.tools and name it nt
### BEGIN SOLUTION
import math as m
import numpy as np
import nose.tools as nt
### END SOLUTION

my_sqrt = m.sqrt(m.pi)
my_other_sqrt = np.sqrt(np.pi)

# check the result to be (almost) equal ...
nt.assert_almost_equal(my_sqrt, my_other_sqrt)

```

<!-- grade_id: at31, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
import math
assert_true( m.pi is math.pi  and  m.sqrt is math.sqrt )
### END HIDDEN TESTS
```

<!-- grade_id: at32, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_almost_equal
import math
assert_almost_equal(my_sqrt, math.sqrt(math.pi))
### END HIDDEN TESTS
```

<!-- grade_id: at33, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_almost_equal
assert_almost_equal(my_other_sqrt, np.sqrt(np.pi))
### END HIDDEN TESTS
```

<!-- grade_id: at34, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# should work without an additional
from nose.tools import assert_almost_equal
assert_almost_equal(my_sqrt, my_other_sqrt)
### END HIDDEN TESTS
```

#### Exercise 4: the below code should print: `Every word is now in the correct order`. Initially, however, when you first run it, it prints: `Every sentence is now in the wrong order`. You are asked, by moving around whole lines of code, to rearrange the code such that it produces the desired output. 

You should:

 1.   not add any code
 2.   not take out (delete) any code
 3.   not split up lines, change, or move around parts of one line

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
# print the right answer with the code elements given below
# 
# the right answer to this exercise is:
#
# "Every word is now in the correct order" 
#
# If your code prints this, your done.

word = 'sentence'

yes_or_no = True
yes_or_no = False

def right_answer(flag):
    if flag == True:
        return 'the correct order'
    else:
        return 'the wrong order'

ans1 = 'Every {0} is now in '.format(word)
ans2 = right_answer(yes_or_no)

def do_print(ans1, ans2):
    print(ans1+ans2)

do_print(ans1,ans2)

word = 'word'

### BEGIN SOLUTION

# print the right answer with the code elements given below
# 
# the right answer to this exercise is:
#
# "Every word is now in the correct order" 
#
# If your code prints this, your done.

def do_print(ans1, ans2):
    print(ans1+ans2)

def right_answer(flag):
    if flag == True:
        return 'the correct order'
    else:
        return 'the wrong order'

word = 'sentence'
word = 'word'

yes_or_no = False
yes_or_no = True

ans1 = 'Every {0} is now in '.format(word) 
ans2 = right_answer(yes_or_no)

do_print(ans1, ans2 )

### END SOLUTION
```

Output:
```
Every sentence is now in the wrong order
Every word is now in the correct order
```

<!-- grade_id: at41, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(yes_or_no, True)
### END HIDDEN TESTS
```

<!-- grade_id: at42, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(word, 'word')
### END HIDDEN TESTS
```

<!-- grade_id: at43, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(ans1, 'Every word is now in ')
### END HIDDEN TESTS
```

<!-- grade_id: at44, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(ans2, 'the correct order')
### END HIDDEN TESTS
```

<!-- grade_id: at45, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):

    def test_execution_flow_revision(self):
        
        expected_output = 'Every word is now in the correct order'
        
        with captured_output() as (out, err):
            do_print(ans1, ans2)
    
        # This can go inside or outside the `with` block
        output = out.getvalue().strip()
        self.assertEqual(output, expected_output)
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_execution_flow_revision (__main__.MyEntryTestClasses) ... ok

----------------------------------------------------------------------
Ran 1 test in 0.002s

OK
```

#### Exercise 5: given a list of samples `all_samples`, containing all samples taken from some process. Samples have a number (an *id*) representing their time of sampling: (example: a sample `1630` has been taken at `16:30h`). This list `all_samples` needs to be filtered, as follows:

- remove all rejected samples (samples to reject have been specified by their id's in an auxiliary list called `samples_rejected`)
- remove all samples taken after some deadline (recall that a sample id reflects the time of sampling)
 
Below is a skeleton, containing a couple of functions you are to complete. Using these functions, you are asked to split `all_samples` in two new lists:

 1. list `regular_samples` of all samples that are **neither rejected, nor late**
 2. a list `samples_removed` of **all other samples**, i.e. removed from the original list
 
 Take 20:00h as the deadline (i.e., remove all samples > 2000)

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
# the list of all samples ...
all_samples       = [  155,  230,  340,  352,  455,  507,  710,  830,  900, 
                      1005, 1120, 1121, 1410, 1530, 1630, 1810, 1930, 2021, 
                      2043, 2136, 2151, 2201 ]

# the list of samples rejected ...
samples_rejected  = [  340, 352, 507, 1120,  1121, 1930 ]

def sample_rejected(s):
    '''return T|F s in list samples_rejected'''
    ### BEGIN SOLUTION
    return s in samples_rejected
    ### END SOLUTION
    
def late_sample(s, deadline):
    '''return T|F s taken after deadline'''
    ### BEGIN SOLUTION
    return s > deadline
    ### END SOLUTION

def regular_sample(s, deadline):
    '''return T|F s is regular (i.e., not rejected or late)'''
    ### BEGIN SOLUTION
    return not( sample_rejected(s) or late_sample(s, deadline) )
    ### END SOLUTION

# create the lists and split ...
regular_samples = list()
samples_removed = list()

# deadline for late samples ...
deadline = 2000

for sample in all_samples:
    ### BEGIN SOLUTION
    if not regular_sample(sample, deadline):
        samples_removed.append(sample)
    else:
        regular_samples.append(sample)
    ### END SOLUTIONS

print('regular samples: {:s}'.format(str(regular_samples)))
print('samples removed: {:s}'.format(str(samples_removed)))

```

Output:
```
regular samples: [155, 230, 455, 710, 830, 900, 1005, 1410, 1530, 1630, 1810]
samples removed: [340, 352, 507, 1120, 1121, 1930, 2021, 2043, 2136, 2151, 2201]
```

<!-- grade_id: at51, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# test predicate sample_reject() ...
from nose.tools import assert_true
for sample in all_samples:
    if sample in samples_rejected:
        assert_true( sample_rejected(sample) == True )
    else:
        assert_true( sample_rejected(sample) == False )
### END HIDDEN TESTS
```

<!-- grade_id: at52, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# test predicate late_sample() ...
from nose.tools import assert_true
deadline = 2000
for sample in all_samples:
    assert_true( late_sample(sample,deadline) == (sample > deadline) )
### END HIDDEN TESTS
```

<!-- grade_id: at53, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# test predicate regulat_sample() ...
from nose.tools import assert_true
deadline = 2000
for sample in all_samples:
    reject = sample_rejected(sample)
    late   = late_sample(sample, deadline)
    assert_true( regular_sample(sample, deadline) == (not(reject or late)) )
### END HIDDEN TESTS
```

<!-- grade_id: at54, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
deadline = 2000
late = [s for s in all_samples if s > deadline]
assert_equal(len(samples_removed), len(late)+len(samples_rejected))
### END HIDDEN TESTS
```

<!-- grade_id: at55, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(regular_samples, [155, 230, 455, 710, 830, 900, 1005, 1410, 1530, 1630, 1810])
assert_equal(samples_removed, [340, 352, 507, 1120, 1121, 1930, 2021, 2043, 2136, 2151, 2201])
## END HIDDEN TESTS
```

# End of the Test