# Python Entry Test

## Part 2 of 4: Control flow (for loops, while loops and if-then-else statements)

## Instructions

Below are 5 Exercises: python problems you need to solve, in order to pass this test. For each correct answer, you will be assigned points, with a maximum of 100 points in total. If you score 80 points or higher, you have passed the test. You cannot see your score immediately; your test will be graded as part of a batch job, later on.

Note the following:
1. You are asked to code in **Python3** (version 3), not Python2
2. For some answers, you obtain a *skeleton*, to frame your answer. You are asked to complete it
3. The answer you are expected to given starts where it says:
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```
   Start your answer by *removing these two lines*. Do not put your answer anywhere else than *just here, in this box*
4. You may use as many lines for your answer as you see necessary
5. In general, however, no more than 5 lines of code are needed
6. Ignore all boxes that are completely empty; just focus on `# YOUR CODE HERE`

## Begin of the Test

#### Exercise 1: complete the following function, that is to print all integer numbers $m \in [0..10]$, using a for loop. The expected output is: 
```
integer number: 0
integer number: 1
integer number: 2
...
integer number: 10
```

11 lines in total, no tabs before, just lowercase, and a single space after the colon `:`

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
def print_numbers_in_for_loop(first, last):
    # for all integer numbers m in [first..last] ...
    ### BEGIN SOLUTION
    for m in range(first, last+1):
        print('integer number: {:d}'. format(m))
    ### END SOLUTION

print_numbers_in_for_loop(0,10)
```

Output:
```
integer number: 0
integer number: 1
integer number: 2
integer number: 3
integer number: 4
integer number: 5
integer number: 6
integer number: 7
integer number: 8
integer number: 9
integer number: 10
```

<!-- grade_id: at1, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):

    def test_q1_print_numbers_for_loop(self):
        
        with captured_output() as (out, err):
            print_numbers_in_for_loop(0, 10)
    
        # This can go inside or outside the `with` block
        output = out.getvalue().strip()
        self.assertEqual(output, 
        'integer number: 0\n'
        'integer number: 1\n'
        'integer number: 2\n'
        'integer number: 3\n'
        'integer number: 4\n'
        'integer number: 5\n'
        'integer number: 6\n'
        'integer number: 7\n'
        'integer number: 8\n'
        'integer number: 9\n'
        'integer number: 10')
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_q1_print_numbers_for_loop (__main__.MyEntryTestClasses) ... ok

----------------------------------------------------------------------
Ran 1 test in 0.002s

OK
```

#### Exercise 2: given a sorted list `my_list = [0,1,2,3,4,5,6,7,8,9,10]`. Also given: 4 initially empty lists `list_1, list_2, list_3, list_4` and 4 half-open intervals:

\begin{enumerate}
\item \hspace{1.0cm} $I_1 = [0..2)$

\item \hspace{1.0cm} $I_2 = [2..5)$   

\item \hspace{1.0cm} $I_3 = [5..8)$   

\item \hspace{1.0cm} $I_4 = [8..11)$   
\end{enumerate}

Copy each of the integer numbers in `my_list` belonging to interval $I_k$ to `list_k`, using `if-elif-else` statements. To copy a number `m` to `list_k` use: `list_k.append(m)`.

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
my_list = [0,1,2,3,4,5,6,7,8,9,10]
list_1, list_2, list_3, list_4 = [], [], [], []

# sub-optimal but most likely solution ...
for m in my_list:
    # put each m in the correct list ...
    ### BEGIN SOLUTION
    if m < 2:
        list_1.append(m)
    elif m >= 2 and m < 5:
        list_2.append(m)
    elif m >= 5 and m < 8:
        list_3.append(m)
    else:
        list_4.append(m)
    ### END SOLUTION
    
# print all lists...
for k, lst in enumerate([list_1, list_2, list_3, list_4]):
    print('list {:d}: {:s}'.format(k+1, str(lst)))
```

Output:
```
list 1: [0, 1]
list 2: [2, 3, 4]
list 3: [5, 6, 7]
list 4: [8, 9, 10]
```

<!-- grade_id: at21, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_1),2)
### END HIDDEN TESTS
```

<!-- grade_id: at22, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_1), [0, 1])
### END HIDDEN TESTS
```

<!-- grade_id: at23, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_2),3)
### END HIDDEN TESTS
```

<!-- grade_id: at24, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_2), [2, 3, 4])
### END HIDDEN TESTS
```

<!-- grade_id: at25, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_3),3)
### END HIDDEN TESTS
```

<!-- grade_id: at26, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_3), [5, 6, 7])
### END HIDDEN TESTS
```

<!-- grade_id: at27, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_4),3)
### END HIDDEN TESTS
```

<!-- grade_id: at28, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_4), [8, 9, 10])
### END HIDDEN TESTS
```

#### Exercise 3
Let $P > 0$ be an arbitrary big number. Recall that for a big enough $n$, we have that $P^{\left( \frac{1}{2} \right) ^n} \approx 1$. 

Repeatedly take the square root of $P$ until:

$|~P-1~| \leq \epsilon$

Here, $|P-1|$ denotes the absolute value of $P-1$. Use $\epsilon = 1.0E-06$. 

Complete the below skeleton.

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
import math

bigNum = 879857982742398758947958728934798758934

EPS = 1.0e-06
number = bigNum

n_root = 0
# while number more than EPS away from 1.0 do ...
### BEGIN SOLUTION
while abs(number -1.0) > EPS:   
### END SOLUTION
    number = math.sqrt(number)
    n_root += 1
    if n_root >= 50:
        break

    
print('the {:d}-th root of {:d} = 1.000000 +/- {:8.6e}'. format(n_root, bigNum, EPS))
```

Output:
```
the 27-th root of 879857982742398758947958728934798758934 = 1.000000 +/- 1.000000e-06
```

<!-- grade_id: at31, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# is the while loop broken?
assert n_root < 50
### END HIDDEN TESTS
```

<!-- grade_id: at32, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(EPS,1.0e-06)
### END HIDDEN TESTS
```

<!-- grade_id: at33, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_greater
assert_greater(bigNum, 0.0)
### END HIDDEN TESTS
```

<!-- grade_id: at34, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_greater_equal
assert_greater_equal (EPS, abs(number-1.0))
### END HIDDEN TESTS
```

<!-- grade_id: at35, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# check whether big number and number are on the same side of 1
from nose.tools import assert_true
assert_true ( (bigNum > 1.0 and number > 1.0) or (bigNum < 1.0 and number < 1.0) )
### END HIDDEN TESTS
```

#### Exercise 4: let `my_list` be a list of items. Remove items from that list in a while loop, one-by-one, until you find that the list is empty. Print how many items the list apparently had. Do not use function `len()`; use  `item = my_list.pop()` to remove an item from the list, and increment the item count. Avoid popping from an empty list.

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
my_list = ['a', 1, 5, '6', (3, 5), 13, 11, 'Z', '!', 5, 2, 4, 9, 6, 'y']

item_count = 0
# while my_list is not empty do ...
### BEGIN SOLUTION
while my_list:
### END SOLUTION
    item = my_list.pop()
    item_count += 1
    if item_count >= 50:
        break

print('my_list contained {:d} items'.format(item_count))
```

Output:
```
my_list contained 15 items
```

<!-- grade_id: at41, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert len(my_list) == 0 # really empty?
### END HIDDEN TESTS
```

<!-- grade_id: at42, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( item_count < 50 ) # loop not broken ?
### END HIDDEN TESTS
```

<!-- grade_id: at43, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(item_count,15)
### END HIDDEN TESTS
```

#### Exercise 5: Given a list `points` of 10 points $(x,y)$ in the plane. We select all points from `points` lieing in the first quadrant, i.e., above the $X$-axis ( line $y=0$ ) and right of the $Y$-axis ( line $x=0$ ) and collect them in a list `Q1`; the code skeleton below shows this. You are asked to select all other points (i.e., points lieing in quadrant two, three and four) by adding a second for-loop, with inverted  selection criteria, and collect these points in a single second list `Q234`. Print the list `Q234` as a string.

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
X0, Y0 = 0.0, 0.0

points = [(1,3), (-1,-2), (1,-1), (-2,0), (3,-1), (-4,1), (1,2), (-1,-3), (2,3), (-4,-1)]

Q1 = []
# make a for-loop to find points in quadrant 1 ...
for p in points:
    if p[0] > X0 and p[1] > Y0:
        Q1.append(p)
    
print('{:d} points in Q1  : {:s}'.format(len(Q1), str(Q1)))

Q234 = []
# now make a 2nd for-loop and invert the selection to find points in quadrant 2 to 4 ...
for p in points:
    ### BEGIN SOLUTION
    if p[0] < X0 or p[1] < Y0:
        Q234.append(p)
    ### END SOLUTION
    
# print the Q234 list ...
### BEGIN SOLUTION
print('{:d} points in Q234: {:s}'.format(len(Q234), str(Q234)))
### END SOLUTION
```

Output:
```
3 points in Q1  : [(1, 3), (1, 2), (2, 3)]
7 points in Q234: [(-1, -2), (1, -1), (-2, 0), (3, -1), (-4, 1), (-1, -3), (-4, -1)]
```

<!-- grade_id: at51, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
# check lieing in quadrant 1
for p in Q1:
    assert_true( p[0] > 0 and p[1] > 0 )
### END HIDDEN TESTS
```

<!-- grade_id: at52, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( len(Q1) == 3 )
### END HIDDEN TESTS
```

<!-- grade_id: at53, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
X0, Y0 = 0, 0
# check lieing in quadrant 2, 3, or 4
for p in Q234:
    assert_true( p[0] < 0 or p[1] < 0 )
### END HIDDEN TESTS
```

<!-- grade_id: at54, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( len(Q234) == len(points) - len(Q1) )
### END HIDDEN TESTS
```

## End of the Test