# Python Entry Test¶

## Part 4 of 4: Basic plotting

## Instructions

Below are 5 Exercises: python problems you need to solve, in order to pass this test. For each correct answer, you will be assigned points, with a maximum of 100 points in total. If you score 80 points or higher, you have passed the test. You cannot see your score immediately; your test will be graded as part of a batch job, later on.

Note the following:
1. You are asked to code in **Python3** (version 3), not Python2
2. For some answers, you obtain a *skeleton*, to frame your answer. You are asked to complete it
3. The answer you are expected to given starts where it says:
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```
   Start your answer by *removing these two lines*. Do not put your answer anywhere else than *just here, in this box*
4. You may use as many lines for your answer as you see necessary
5. In general, however, no more than 5 lines of code are needed
6. Ignore all boxes that are completely empty; just focus on `# YOUR CODE HERE`

# Begin of the Test

#### Exercise 1: make a plot of a given function

Given a function: $y = (x-5)^2 + 2$. 

 1. prepare the function coordinates: $(x, F(x))$ in two lists `X` and `Y` in the below code skeleton
 
 2. plot the function as a red solid line (use: `plt.plot(..., ..., 'r-')`), using the below given code skeleton
 
 3. plot a horizontal asymptote (line) $y = constant = y_{min}$ and a vertical symmetry line $x=constant$ as dashed blue lines. 
 
    For 3. do this:   
 
    * Create lists `X_asymp` and `Y_asymp` and `X_symm` and `Y_symm` to contain the coordinates for the asymptote and the vertical symmetry line, resp.    
 
    * to plot a straight line, we need at least 2 points (begin and end point). Compute or simply code in the below skeleton two points $(x_0, y_{min})$ and $(x_1, y_{min})$ for the horizontal asymptote. Append these two points to the lists: `X_asymp` and `Y_asymp`, resp. (remember we need separate lists for X and Y)   
 
    * do the same for the symmetry line   
 
    * then: plot each of them in the same way as the plotting in 2. To obtain a dashed blue line, use: `plt.plot(..., ..., 'b--')`

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
%matplotlib inline

# use pyplot ...
import matplotlib.pyplot as plt

# program your function to plot below ...
def F(x):
    # specify y=F(x) below ...
    # like in: y = (x ...)**2 etc.
    # YOUR CODE HERE
    raise NotImplementedError()
    return y

# nr of points acros X-domain ...
N = 10+1 # 1 extra for endpoint ...

# empty lists to hold your function x and y coordinates ...
X, Y = [], []

# use this code to generate your data ...
for x in range(N):
    # append the next x to X ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # ... and the F(x) to Y ...
    # YOUR CODE HERE
    raise NotImplementedError()

# plot the line using the prepared lists X and Y ...

# plt.plot( list of x-coordinates, list of y-coordinates, 'r-')
# YOUR CODE HERE
raise NotImplementedError()

# do something similar here, for the y_min line and line x=5 ...
X_asymp, Y_asymp = [], []
X_symm,  Y_symm  = [], []

# APPEND COORDINATES needed to plot the
# asymptote and the symmetry line BELOW
# like so:
#    x_0 = ...
#    x_1 = ...
#    X_asymp.append(x_0)
#    X_asymp.append(x_1)
#    ...
# then issue a plot comand to plot the lines ...

# YOUR CODE HERE
raise NotImplementedError()

# ----- do not change below -------
plt.grid(True)
plt.title('$F(x)=(x-5)^{2}+2$,   for $x \in [0..10]$')
plt.show()
```

Output:
```
~/.local/lib/python3.8/site-packages/IPython/core/interactiveshell.py in enable_matplotlib(self, gui)
   3515         """
   3516         from IPython.core import pylabtools as pt
-> 3517         from matplotlib_inline.backend_inline import configure_inline_support
   3518         gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3519 

~/.local/lib/python3.8/site-packages/matplotlib_inline/backend_inline.py in <module>
      4 # Distributed under the terms of the BSD 3-Clause License.
      5 
----> 6 import matplotlib
      7 from matplotlib.backends.backend_agg import (  # noqa
      8     new_figure_manager,

ModuleNotFoundError: No module named 'matplotlib'
```

<!-- grade_id: at11, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        N_test = 10+1
        self.X_test = [x for x in range(N_test)]
        self.Y_test = [(x-5)**2 + 2 for x in range(N_test)]
        
        
    def test_F_producing_correct_y(self):
        for k in range(N):
            self.assertEqual(self.X_test[k], X[k])
            self.assertEqual(self.Y_test[k], F(X[k]))
            
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_F_producing_correct_y (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_F_producing_correct_y (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-2-ac23643f70db>", line 13, in test_F_producing_correct_y
    for k in range(N):
NameError: name 'N' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.003s

FAILED (errors=1)
```

<!-- grade_id: at12, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):  
        
    def setUp(self):
        self.N_test = 10+1
        self.X_test = [x for x in range(self.N_test)]
        self.Y_test = [(x-5)**2 + 2 for x in range(self.N_test)]
        
    def test_data_lists_dimensions_correct(self):
        self.assertEqual(self.N_test, N)
        self.assertEqual(len(self.X_test), len(X))
        self.assertEqual(len(self.Y_test), len(Y))

    def tearDown(self):
        pass
    
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_data_lists_dimensions_correct (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_data_lists_dimensions_correct (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-3-b512f94d975d>", line 12, in test_data_lists_dimensions_correct
    self.assertEqual(self.N_test, N)
NameError: name 'N' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.002s

FAILED (errors=1)
```

<!-- grade_id: at13, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.X_asymp_test = [0, 10]
        self.Y_asymp_test = [2, 2]
        
    def test_plot_correct_asympthotics(self):
        #self.assertEqual(self.X_asymp_test, X_asymp)
        self.assertEqual(self.Y_asymp_test, Y_asymp)
            
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_plot_correct_asympthotics (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_plot_correct_asympthotics (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-4-977f055a246b>", line 12, in test_plot_correct_asympthotics
    self.assertEqual(self.Y_asymp_test, Y_asymp)
NameError: name 'Y_asymp' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.002s

FAILED (errors=1)
```

<!-- grade_id: at14, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.X_symm_test = [5, 5]
        self.Y_symm_test = [0, 25+2]
        
    def test_plot_correct_asympthotics(self):
        self.assertEqual(self.X_symm_test, X_symm)
        # self.assertEqual(self.Y_symm_test, Y_symm)
            
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_plot_correct_asympthotics (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_plot_correct_asympthotics (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-5-aaf6f7ce591e>", line 11, in test_plot_correct_asympthotics
    self.assertEqual(self.X_symm_test, X_symm)
NameError: name 'X_symm' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.003s

FAILED (errors=1)
```

<!-- grade_id: at15, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        N_test = 10+1
        self.X_test = [x for x in range(N_test)]
        self.Y_test = [(x-5)**2 + 2 for x in range(N_test)]
        
    def test_plot_called_with_correct_data(self):
        pass
            
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_plot_called_with_correct_data (__main__.MyEntryTestClasses) ... ok

----------------------------------------------------------------------
Ran 1 test in 0.002s

OK
```

#### Exercise 2: make a scatter plot of 40 randomly distributed points (x,y)∈[0,1)×[0,1)

Finalize the code skeleton given below to produce the following plot:

![expected result 2](./figures/expected_result2.png)

 1. specify the domains
 2. specify the correct number of point N to generate
 3. generate N random point coordinates $(x,y)$ and append them to two lists `X` and `Y`
 4. specify a blue small dot marker string
 5. plot the points

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
%matplotlib inline

# import random generator ...
import random as rnd

# define the x and y domain ...
#XMIN, XMAX = ( , )
#YMIN, YMAX = ( , )
# YOUR CODE HERE
raise NotImplementedError()

# how many points to generate?
# specify N below ...
#N = 
# YOUR CODE HERE
raise NotImplementedError()

# create empty coordinates lists...
X, Y = [], []

# generate N random data point coordinates ...
# and append them to X and Y ...
for n in range(N):
    # rnd.random() returns random float in [0,1)
    # ... generate a random x coordinate and an y ...
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # ... add them to the list to plot ...
    
    # YOUR CODE HERE
    raise NotImplementedError()

# add the marker string to plot a blue dot ...
#marker = 
# YOUR CODE HERE
raise NotImplementedError()

# uncomment the following line and fill in X, Y and marker ...
# plt.plot( list of x-coordinates, list of y-coordinates, marker string)
# YOUR CODE HERE
raise NotImplementedError()

# ----- do not change below -------
plt.grid(True)
title = 'scatterplot random points on [{:4.1f},{:4.1f}] x [{:4.1f},{:4.1f}] (N={:d})'. \
          format(XMIN,XMAX,YMIN,YMAX, N)
plt.title(title)
plt.show()   
```

Output:
```
~/.local/lib/python3.8/site-packages/IPython/core/interactiveshell.py in enable_matplotlib(self, gui)
   3515         """
   3516         from IPython.core import pylabtools as pt
-> 3517         from matplotlib_inline.backend_inline import configure_inline_support
   3518         gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3519 

~/.local/lib/python3.8/site-packages/matplotlib_inline/backend_inline.py in <module>
      4 # Distributed under the terms of the BSD 3-Clause License.
      5 
----> 6 import matplotlib
      7 from matplotlib.backends.backend_agg import (  # noqa
      8     new_figure_manager,

ModuleNotFoundError: No module named 'matplotlib'
```

<!-- grade_id: at21, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.XDOMAIN,self.YDOMAIN = (0, 1), (0, 1)
              
    def test_domain_dimensions(self):
        '''verify domain parameters are chosen correctly'''
        self.assertEqual(XMIN, min(self.XDOMAIN))
        self.assertEqual(XMAX, max(self.XDOMAIN))
        self.assertEqual(YMIN, min(self.YDOMAIN))
        self.assertEqual(YMAX, max(self.YDOMAIN))
        
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_domain_dimensions (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly ... ERROR

======================================================================
ERROR: test_domain_dimensions (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-8-e48d16d905e0>", line 11, in test_domain_dimensions
    self.assertEqual(XMIN, min(self.XDOMAIN))
NameError: name 'XMIN' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at22, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
        
    def test_correct_N(self):
        '''verify N correctly specified'''
        
        expected_N = 40
        self.assertEqual(N, expected_N)

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_correct_N (__main__.MyEntryTestClasses)
verify N correctly specified ... ERROR

======================================================================
ERROR: test_correct_N (__main__.MyEntryTestClasses)
verify N correctly specified
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-9-ecfed7d4f6ee>", line 10, in test_correct_N
    self.assertEqual(N, expected_N)
NameError: name 'N' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at24, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.XDOMAIN,self.YDOMAIN = (0, 1), (0, 1)
        
    def test_X_Y_in_domain_boundaries(self):
        '''verify generated data within domain boundaries'''
        
        for n in range(N): 
            self.assertGreater(X[n],   min(self.XDOMAIN))
            self.assertLessEqual(X[n], max(self.XDOMAIN))
            
            self.assertGreater(Y[n],   min(self.YDOMAIN))
            self.assertLessEqual(Y[n], max(self.YDOMAIN))

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_X_Y_in_domain_boundaries (__main__.MyEntryTestClasses)
verify generated data within domain boundaries ... ERROR

======================================================================
ERROR: test_X_Y_in_domain_boundaries (__main__.MyEntryTestClasses)
verify generated data within domain boundaries
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-10-719944a69c68>", line 12, in test_X_Y_in_domain_boundaries
    for n in range(N):
NameError: name 'N' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at23, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
        
    def test_marker(self):
        '''verify marker plots a blue dot symbol'''
        
        self.assertTrue( 'b' in marker )
        self.assertTrue( '.' in marker or '*' in marker or 'o' in marker)
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_marker (__main__.MyEntryTestClasses)
verify marker plots a blue dot symbol ... ERROR

======================================================================
ERROR: test_marker (__main__.MyEntryTestClasses)
verify marker plots a blue dot symbol
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-11-3f82cea27d44>", line 9, in test_marker
    self.assertTrue( 'b' in marker )
NameError: name 'marker' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

#### Exercise 3: plot random points in a scatter plot and recolor points belonging to a given sub-domain

Random points $(u,v)$ are randomly generated by a random generator on a unit domain $[0,1)$. You are to map and plot 100 random points so generated, on the wider target domain $[X_{min},X_{max}) \times [Y_{min}, Y_{max})$. To accomplish this, we are going to define and program a *shift* and a *scaling*, with which we can *map* the unit domain on the target domain. See the below sketch.

 To map $(u,v) \rightarrow (x,y)$ use:
```python
    (x,y)=(shift + scale*u, shift + scale*v)
```
in which:
- `shift` is $O_d- O_r = -10$, and
- `scale` is $X~range_d / range_r = 20$.

<img width="380" src="./figures/transform.png">

We are going to do the above pltting task step-by-step. And to help you, we provide you again with a code skeleton. 

Do this:

1. implement the `shift` and `scale` as described above, in the below code skeleton

2. implement the generation of unit random numbers `(u,v)`, and its mapping onto `(x,y)` using the `shift` and `scale` factors computed under 1. Add the so mapped points `(x,y)` to the lists `X` and `Y`

3. plot points $(x,y)$ that are **in quadrant 1** (i.e, the points for which: $x > 0 \wedge y > 0$) using **red square marker**, and points in the other quadrants (for which either $x \leq 0 \vee y \leq 0$), with a **blue dot marker**. Do the plotting in the **left subplot** (the left part of the plot).

Additional hints for 3: 

* for plotting a red square, you code in the below code skeleton: `marker = 'rs'`. For a blue dot, set: `marker = 'bo'`
* to plot in the left subplot, you use the code: `ax[0].plot()` in the code skeleton
* to have an impression of what is asked from you, see the expected plot given in the below figure

4. in the right subplot, we are going to plot the same data, but with their markers *reversed*: the blue dot marker for quadrant 1 and the red square marker for every point not in quadrant 1. We could just change the markers, or leave the markers as is, and change the selection of points (the line: `p > x0 ...` in your code). Do the latter: update the selection of the points to be plotted **now** with a red square, the rest with a blue dot

Additional hints for 4:

* what you actually are asked to do, is to select the **logical complement** of the data points that you plotted with a red marker in 3. You **logically invert** the selection
* now use `ax[1].plot()` to make your code plotting in the right subplot

The expected result is shown below.

![expected result](./figures/expected_result3.png)

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
import matplotlib.pyplot as plt

import numpy


# domain (equal for x and y)
target_domain  = (-10,10)
unit_domain    = (  0, 1)

# indices into domain
MIN, MAX =    0, 1

N = 100

# using the above explanation on `shift` and `scale`, program
# below the computation of these two transform-factors, like so:
#shift   = target_domain[...] - unit_domain[...]
#range_d = target_domain[...] - target_domain[...]
#range_r = unit_domain[...]   - unit_domain[...]
#scale = ... / ...
# YOUR CODE HERE
raise NotImplementedError()

X, Y = numpy.empty(N), numpy.empty(N)
# below, you are going to use scale and shift to
# compute X and Y coordinates of the generated points,
# like so:
#for j in range(N):
#    u = numpy.random.random()
#    v = numpy.random.random()
#    X[j] = shift + scale*u
#    Y[j] = shift + scale*v
# YOUR CODE HERE
raise NotImplementedError()

NR_ROWS, NR_COLS, FIG_SIZE = 1, 2, (10,4)
fig, ax = plt.subplots(NR_ROWS, NR_COLS, figsize=FIG_SIZE)
fig.tight_layout(pad=2.0, w_pad=4.5, h_pad=0.0)

# use: ax[0] to plot in the left subplot, ax[1] for the right subplot 

ax[0].set_xlim(target_domain)
ax[0].set_ylim(target_domain)
# we compute x0 and y0 as the central x and y in the target_domain
# this is where we put the axis system ...
x0 = y0 = ( target_domain[MIN] + target_domain[MAX]) / 2
ax[0].plot([target_domain[MIN], target_domain[MAX]], [y0, y0], 'k--', linewidth=2.0)
ax[0].plot([x0, x0], [target_domain[MIN], target_domain[MAX]], 'k--', linewidth=2.0)

for p, q in zip(X,Y):
    # with next p from X and q from Y ...
    # ... test in which quadrant ...
    # ... plot with the correct marker ...
    # ... set the marker to:  marker = 'rs'   for a red square
    # ....or:            to:  marker = 'bo'   for a blue dot

    # YOUR CODE HERE
    raise NotImplementedError()
    
ax[0].grid(axis='both', which='major')
ax[0].set_title('{:d} samples, red: quadrant 1'.format(N))
    
# next, in the right subplot, plot the inverted data ...
ax[1].set_xlim(target_domain)
ax[1].set_ylim(target_domain)
ax[1].plot([target_domain[MIN], target_domain[MAX]], [y0, y0], 'k--', linewidth=2.0)
ax[1].plot([x0, x0], [target_domain[MIN], target_domain[MAX]], 'k--', linewidth=2.0)

# now invert your selection of red squared points ...

for p, q in zip(X,Y):
    # YOUR CODE HERE
    raise NotImplementedError()
    
ax[1].grid(axis='both', which='major')
ax[1].set_title('{:d} samples, red: quadrant 2,3,4'.format(N))

plt.show()
```

Output:
```
~/.local/lib/python3.8/site-packages/IPython/core/interactiveshell.py in enable_matplotlib(self, gui)
   3515         """
   3516         from IPython.core import pylabtools as pt
-> 3517         from matplotlib_inline.backend_inline import configure_inline_support
   3518         gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3519 

~/.local/lib/python3.8/site-packages/matplotlib_inline/backend_inline.py in <module>
      4 # Distributed under the terms of the BSD 3-Clause License.
      5 
----> 6 import matplotlib
      7 from matplotlib.backends.backend_agg import (  # noqa
      8     new_figure_manager,

ModuleNotFoundError: No module named 'matplotlib'
```

<!-- grade_id: at31, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.XDOMAIN,self.YDOMAIN = (-10, 10), (-10, 10)
        self.X_RANGE = max(self.XDOMAIN) - min(self.XDOMAIN)
        self.Y_RANGE = max(self.YDOMAIN) - min(self.YDOMAIN)
              
    def test_transform_parms(self):
        '''verify transform parameters set correctly'''
        
        self.assertEqual(shift, target_domain[MIN]-unit_domain[MIN])
        self.assertEqual(range_d, target_domain[MAX]-target_domain[MIN])
        self.assertEqual(range_d, self.X_RANGE)
        self.assertEqual(range_r, unit_domain[MAX]-unit_domain[MIN])
        self.assertEqual(range_r, 1)
        self.assertEqual(scale, range_d/range_r)
        
    def tearDown(self):
        pass
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_transform_parms (__main__.MyEntryTestClasses)
verify transform parameters set correctly ... ERROR

======================================================================
ERROR: test_transform_parms (__main__.MyEntryTestClasses)
verify transform parameters set correctly
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-13-314934d377ab>", line 14, in test_transform_parms
    self.assertEqual(shift, target_domain[MIN]-unit_domain[MIN])
NameError: name 'shift' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at32, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.DOMAIN = (-10, 10)
        self.RANGE = max(self.DOMAIN) - min(self.DOMAIN)
              
    def test_transformed_in_range_domain(self):
        '''verify all points lie in the correct target domain'''
        for (x,y) in zip(X,Y):
            self.assertGreaterEqual(x, min(self.DOMAIN))
            self.assertGreaterEqual(y, min(self.DOMAIN))
            self.assertLess(x,min(self.DOMAIN)+self.RANGE)
            self.assertLess(y,min(self.DOMAIN)+self.RANGE)
            
    def test_transform(self):
        expected_shift = -10.0
        expected_scale = +20.0
        self.assertAlmostEqual(shift, expected_shift)
        self.assertAlmostEqual(scale, expected_scale)
    
    def test_transformation(self):
        test_u,test_v = 0.25,  0.40
        self.assertAlmostEqual(shift + scale * test_u, -5.0, 6)
        self.assertAlmostEqual(shift + scale * test_v, -2.0, 6)
        
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_transform (__main__.MyEntryTestClasses) ... ERROR
test_transformation (__main__.MyEntryTestClasses) ... ERROR
test_transformed_in_range_domain (__main__.MyEntryTestClasses)
verify all points lie in the correct target domain ... ERROR

======================================================================
ERROR: test_transform (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-14-3065d4c5d98a>", line 21, in test_transform
    self.assertAlmostEqual(shift, expected_shift)
NameError: name 'shift' is not defined

======================================================================
ERROR: test_transformation (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-14-3065d4c5d98a>", line 26, in test_transformation
    self.assertAlmostEqual(shift + scale * test_u, -5.0, 6)
Name
```

<!-- grade_id: at33, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.XDOMAIN,self.YDOMAIN = (-10, 10), (-10, 10)
    
    def test_points_mapped(self):
        MIN, MAX = (0,1)
        for p,q in zip(X,Y):
            self.assertTrue( p >= self.XDOMAIN[MIN] )
            self.assertTrue( p <  self.XDOMAIN[MAX] )
            self.assertTrue( q >= self.YDOMAIN[MIN] )
            self.assertTrue( q <  self.YDOMAIN[MAX] )            
            
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_points_mapped (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_points_mapped (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-15-5ecb374a9e20>", line 11, in test_points_mapped
    for p,q in zip(X,Y):
NameError: name 'X' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at34, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.XDOMAIN,self.YDOMAIN = (-10, 10), (-10, 10)
              
    def test_X_Y_in_domain(self):
        '''verify generate points in domain'''
        
        # test the last combinatio used ...
        if p <= 0 or q <= 0:
            self.assertEqual(marker, 'rs')
        else:
            self.assertEqual(marker, 'bo')
        
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_X_Y_in_domain (__main__.MyEntryTestClasses)
verify generate points in domain ... ERROR

======================================================================
ERROR: test_X_Y_in_domain (__main__.MyEntryTestClasses)
verify generate points in domain
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-16-c01e4d57e150>", line 13, in test_X_Y_in_domain
    if p <= 0 or q <= 0:
NameError: name 'p' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

#### Exercise 4: plot two orthogonal vectors and a circle, using a given code skeleton

Given a vector $\vec{v}=(4,2)$, starting from origin $O=(0,0)$ and heading to point $P=(4,2)$. We want to create a vector $\vec{w}=(-v[1],v[0])$, such that $\vec{w} \perp \vec{v}$. Also, we want to show that the length (the *norm*) $| \vec{w} |~=~ | \vec{v}|$. We achieve this by drawing a circle with radius $R=|\vec{v}|$, showing that both $\vec{v}$ and $\vec{w}$ end exactly on this circle.

Using the below code skeleton, your are asked to:
1. implement `Orig = (0,0)`, `v = (4,2)`, and perpendicular `w` as defined above
2. implement $L = |\vec{v}| = \sqrt{v_x^2 + v_y^2}$, the length of vector `v` which is equal to the distance of point `(4,2)` from the origin in this case
3. run the program to plot the vectors implemented, and the circle with radius $L$

The expected result is shown below.

![expected result 4](./figures/expected_result4.png)


<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
import math
import matplotlib.pyplot as plt

XMIN, XMAX = -10, +10
YMIN, YMAX = -10, +10

# instantiate: Orig, v, w below (each a 2-tuple (x,y)) ...
# Orig is given below, as an example of what is epected from you ...
# Orig = (0,0)   # origin (begin point) of the vectors v and w
# v = ( , )      # destination (end point) of vector v
# w = ( , )      # same, vector w, w perpendicular to v
# YOUR CODE HERE
raise NotImplementedError()

# we need these for plotting ...
v_x, v_y  = v[0]-Orig[0], v[1]-Orig[1]
w_x, w_y  = w[0]-Orig[0], w[1]-Orig[1]

# compute the length L using the definition given ...
#L = math.sqrt( ... )
# YOUR CODE HERE
raise NotImplementedError()

# now w're good to plot ... no changes needed below.

# arrow parameters ...
HD_WIDTH  = .5     # head width
HD_LNGTH  = .5     # head length
FC_COLOR  = 'red'  # face color
EC_COLOR  = 'red'  # edge color

#cricle parameters ...
CL_FILL   = False  # do not fill circle
CL_LCOLOR = 'blue'
CL_LSTYLE = '-.'
CL_LWIDTH = 2

# no need to change anything below ...
ax = plt.axes()
ax.set_xlim(XMIN, XMAX)
ax.set_ylim(YMIN, YMAX)

# plot vector v (just fill in first line of command; keep rest)
ax.arrow(x=Orig[0], y=Orig[1], dx=v_x, dy=v_y, \
    head_width=HD_WIDTH, head_length=HD_LNGTH, \
    length_includes_head=True, \
    fc=FC_COLOR, ec=EC_COLOR)
ax.text(2, 1.2, r'$\vec{v}$')

# plot vector w (just fill in first line of command; keep rest)
ax.arrow(x=Orig[0], y=Orig[1], dx=w_x, dy=w_y, \
    head_width=HD_WIDTH, head_length=HD_LNGTH, \
    length_includes_head=True, \
    fc=FC_COLOR, ec=EC_COLOR)
ax.text(-.8, 2, r'$\vec{w}$')

# label the origin ...
ax.text(-0.9, -0.5, r'$O=(0,0)$')

# plot radius ...
ax.arrow(x=Orig[0], y=Orig[1], dx=L, dy=0, \
    head_width=HD_WIDTH/2, head_length=HD_LNGTH/2, \
    length_includes_head=True, fc='k', ec='k')
ax.text(2.2, .1, 'R={:4.2f}'.format(L))

crcle = plt.Circle(Orig, radius=L, \
    fill=CL_FILL, color=CL_LCOLOR, \
    linestyle=CL_LSTYLE, linewidth=CL_LWIDTH)
ax = plt.gca()
ax.add_patch(crcle)
ax.axis('scaled')

plt.grid(True)
plt.show()
```

Output:
```
~/.local/lib/python3.8/site-packages/IPython/core/interactiveshell.py in enable_matplotlib(self, gui)
   3515         """
   3516         from IPython.core import pylabtools as pt
-> 3517         from matplotlib_inline.backend_inline import configure_inline_support
   3518         gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3519 

~/.local/lib/python3.8/site-packages/matplotlib_inline/backend_inline.py in <module>
      4 # Distributed under the terms of the BSD 3-Clause License.
      5 
----> 6 import matplotlib
      7 from matplotlib.backends.backend_agg import (  # noqa
      8     new_figure_manager,

ModuleNotFoundError: No module named 'matplotlib'
```

<!-- grade_id: at41, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.expected_orig = (0,0)
        self.expected_v = (4,2)
        self.expected_w = (-2,4)
              
    def test_correct_ortho(self):
        '''verify domain parameters are chosen correctly'''
        self.assertEqual(Orig, self.expected_orig)
        self.assertEqual(  v,  self.expected_v)
        self.assertEqual(  w,  self.expected_w)
                
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_correct_ortho (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly ... ERROR

======================================================================
ERROR: test_correct_ortho (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-18-b1cfedabb791>", line 13, in test_correct_ortho
    self.assertEqual(Orig, self.expected_orig)
NameError: name 'Orig' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at42, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
              
    def test_function_correct(self):
        '''verify domain parameters are chosen correctly'''
        v = (4,2)
        expected_norm = math.sqrt( v[0]**2 + v[1]**2 )
        self.assertEqual(L, expected_norm)

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_function_correct (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly ... ERROR

======================================================================
ERROR: test_function_correct (__main__.MyEntryTestClasses)
verify domain parameters are chosen correctly
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-19-aa385b0df2a0>", line 9, in test_function_correct
    expected_norm = math.sqrt( v[0]**2 + v[1]**2 )
NameError: name 'math' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

#### Exercise 5: prepare and plot a simple histogram

Data on ball bearing rolling friction coefficients $C_r$ are read from a test data file `Bearings.dat`. The file contains 109 bearing data records: `id, cr`, where `id` is the bearing id, and `cr` is its rolling friction coefficient. The below code skeleton reads in these data already, you are to program the data analysis.

More in detail, you are asked to analyze this $C_r$ value, and identify the bearings (by their id's) with the minimum, and maximum values for $C_r$. Also, compute the average value, using the below skeleton. Finally, plot a histogram with bars for each of the $C_r$ classes (called: bins), and a line of the average value. The expected diagram is shown below.

![expected histogram](./figures/expected_result5.png)


1. in function `printAnalysis()`, implement the statistical analysis;
  * check out all min, max and avg initializations first
  * using a `for` .. `in` loop:
     - find the min, max `cr` (record their `id`)
     - sum up all `cr` values to compute avg afterwards
  * once the loop is done, divide sum by number of bearings to find average
2. run the program, check the printed analysis and the plot and verify the results. Expected results;
```
    Ball bearing analysis (N=109)
    Count  : 109 data values
    Cr min : 2.1000000000e-04 [id 25]
    Cr max : 9.8100000000e-03 [id 31]
    Cr avg : 4.6193577982e-03
```
3. plot a vertical dashed line for `cr` equal to the average value, much like the plotting of the minimum `cr`. To that end, in `plotStats()`, prepare `X` and `Y` with two coordinate pairs for the vertical line. The minimum `cr` value is contained in `mma` and for `Y`, just take frequencies 0 to 20.  Implement and rerun


<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
%matplotlib inline

import matplotlib.pyplot as plt

# import other modules we need ...
import numpy as np
import scipy.stats as st


def loadBearingDataFromFile(flname,delim=','):
    ''' load bearing data from 'flname'; return N, data'''
    id, cr = np.loadtxt(flname, delimiter=delim, comments='#', \
                        usecols=(0,1), unpack=True)
    N = len(id)
    # print('{:d} data records gelezen van file: {:s}'.format(N,flname))
    data = np.empty(N*2, dtype=float).reshape(N,2)
    data[...,0] = id
    data[...,1] = cr
    return data

def printAnalysis(data,domain):
    '''analyze data and return min,max,avg'''
    # indexes data record layout ...
    ID,CR = (0,1)
    CRMIN,CRMAX  = max(domain), min(domain) #(domain[1],domain[0])
    
    # nr data items in data array ?
    N = len(data)
    
    # init running statistics ...
    crmin, crmax, crsum =(CRMIN,CRMAX,0.0)
    # keep track of items ...
    idmin, idmax = (None,None)
    
    # implement the statistical analysis below;
    # check out the above initializations first
    #
    # using a for .. in loop over data:
    # - find min CR-value and store in crmin (register id in idmin)
    # - find max CR-value and store in crmax (register id in idmax)
    # - sum up all CR in crsum
    # once the loop has terminated, divide crsum by nr of bearings 
    # to find cravg
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # statistical analysis ...
    hdr='Ball bearing analysis (N={:d})'.format(N)   
    print(hdr)
    print('Count  : {:d} data values'.    format(N))
    print('Cr min : {:15.10e} [id {:d}]'. format(crmin,int(idmin)))
    print('Cr max : {:15.10e} [id {:d}]'. format(crmax,int(idmax)))
    print('Cr avg : {:15.10e}'.           format(cravg))
    # return main stats ...
    return (crmin,crmax,cravg)

def plotStats(data,domain,mma):
    '''plot data and results in a histogram'''
    
    # define plotting constants constant ...
    # ... #classes (bins) in histo
    NRCLASSES    = 10
    # ... domain interval sizes (width of the bins)
    LOW,HGH,STEP = (domain[0],domain[1],(domain[1]-domain[0])/NRCLASSES)
    # ... indexes in min,max,avg typle
    MIN,MAX,AVG = (0,1,2)
    
    # skeleton for histogram ....
    plt.subplots(1, figsize=(8,4))
    axis = plt.subplot(111)
    
    axis.set_title('Histogram $C_r$ (N={:d})'.format(len(data)))
    axis.set_xlabel('$C_r$')
    axis.set_ylabel('Frequency')
    axis.set_xticks(np.arange(LOW, HGH+STEP, STEP ))
    axis.set_yticks([0,5,10,15,20])
    
    # ... now start plotting the histogram iself ...
    axis.hist(data[...,1], 
              bins=np.arange(LOW, HGH+STEP, STEP), rwidth=.72, \
              color='orange',histtype='bar')
    
    # ...  vertical line for min value found during analyse ...
    axis.plot([mma[MIN],mma[MIN]],[0, 5], \
              color='red',linestyle='--',linewidth=3.0,\
              label='minimum $C_r$')
    
    # in similar vein, plot vertical line for avg below:
    # 1. use the correct X_avg = [x=avg,x=avg]  , as calculated
    # 2. use Y_avg=[0,20] to plot your x=avg line ...
    # X_avg = [ , ]
    # Y_avg = [ , ]
    # YOUR CODE HERE
    raise NotImplementedError()
    
    axis.plot(X_avg, Y_avg, \
              color='brown',linestyle=':',linewidth=3.0,\
              label='average $C_r$')
    axis.grid(which='major',axis='both',alpha=.75)
    # ... best position for legend in plot ...
    plt.legend(loc='best')
    
    return X_avg, Y_avg

# -----------------------------------------------------------

COMMA       = ','
flname      = './resources/Bearings.dat'
# values of Cr belong to this OPEN domain
CRLOW,CRHGH = (0.000, 0.010)    
DOMAIN      = (CRLOW,CRHGH) # 2-tuple domain boundaries 

# ... load data from file ...
data = loadBearingDataFromFile(flname,COMMA)
# ... analyze. print results ...
mma = printAnalysis(data,DOMAIN)
# ... plot histogram ...
X_avg, Y_avg = plotStats(data,DOMAIN,mma)

```

Output:
```
~/.local/lib/python3.8/site-packages/IPython/core/interactiveshell.py in enable_matplotlib(self, gui)
   3515         """
   3516         from IPython.core import pylabtools as pt
-> 3517         from matplotlib_inline.backend_inline import configure_inline_support
   3518         gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3519 

~/.local/lib/python3.8/site-packages/matplotlib_inline/backend_inline.py in <module>
      4 # Distributed under the terms of the BSD 3-Clause License.
      5 
----> 6 import matplotlib
      7 from matplotlib.backends.backend_agg import (  # noqa
      8     new_figure_manager,

ModuleNotFoundError: No module named 'matplotlib'
```

<!-- grade_id: at51, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):
    
    def setUp(self):
        self.domain = (0.000, 0.010)
        self.flname = './resources/Bearings.dat'
        self.data = loadBearingDataFromFile(flname,delim=',')

    def test_stats_bearing_correct(self):
        
        with captured_output() as (out, err):
            printAnalysis(self.data, self.domain)
    
        # This can go inside or outside the `with` block
        output = out.getvalue().strip()
        self.assertEqual(output,
        'Ball bearing analysis (N=109)\n'
        'Count  : 109 data values\n'
        'Cr min : 2.1000000000e-04 [id 25]\n'
        'Cr max : 9.8100000000e-03 [id 31]\n'
        'Cr avg : 4.6193577982e-03')
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_stats_bearing_correct (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_stats_bearing_correct (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-21-5b89477072a9>", line 23, in setUp
    self.data = loadBearingDataFromFile(flname,delim=',')
NameError: name 'loadBearingDataFromFile' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at52, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
              
    def test_mma_correct(self):
        '''verify statistics calculated correctly'''
        
        MIN,MAX,AVG = (0,1,2)
        expected_mma = (2.1000000000e-04, 9.8100000000e-03, 4.6193577982e-03)
        self.assertAlmostEqual(mma[MIN], expected_mma[MIN], 5)
        self.assertAlmostEqual(mma[MAX], expected_mma[MAX], 5)
        self.assertAlmostEqual(mma[AVG], expected_mma[AVG], 5)      
        
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_mma_correct (__main__.MyEntryTestClasses)
verify statistics calculated correctly ... ERROR

======================================================================
ERROR: test_mma_correct (__main__.MyEntryTestClasses)
verify statistics calculated correctly
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-22-23e6613a2663>", line 11, in test_mma_correct
    self.assertAlmostEqual(mma[MIN], expected_mma[MIN], 5)
NameError: name 'mma' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

<!-- grade_id: at53, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest

class MyEntryTestClasses(unittest.TestCase):
              
    def test_mma_line_correct(self):
        '''verify average line correctly located'''
        
        MIN,MAX,AVG = (0,1,2)
        
        self.assertAlmostEqual(X_avg[0], mma[AVG], 5)
        self.assertAlmostEqual(X_avg[1], mma[AVG], 5)
        self.assertGreaterEqual(Y_avg[0], 0)
        self.assertGreaterEqual(Y_avg[1], 1)
        
        
    def tearDown(self):
        pass

        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_mma_line_correct (__main__.MyEntryTestClasses)
verify average line correctly located ... ERROR

======================================================================
ERROR: test_mma_line_correct (__main__.MyEntryTestClasses)
verify average line correctly located
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-23-808454e82217>", line 11, in test_mma_line_correct
    self.assertAlmostEqual(X_avg[0], mma[AVG], 5)
NameError: name 'X_avg' is not defined

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

## End of the Test