# Python Entry Test

## Part 2 of 4: Control flow (for loops, while loops and if-then-else statements)

## Instructions

Below are 5 Exercises: python problems you need to solve, in order to pass this test. For each correct answer, you will be assigned points, with a maximum of 100 points in total. If you score 80 points or higher, you have passed the test. You cannot see your score immediately; your test will be graded as part of a batch job, later on.

Note the following:
1. You are asked to code in **Python3** (version 3), not Python2
2. For some answers, you obtain a *skeleton*, to frame your answer. You are asked to complete it
3. The answer you are expected to given starts where it says:
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```
   Start your answer by *removing these two lines*. Do not put your answer anywhere else than *just here, in this box*
4. You may use as many lines for your answer as you see necessary
5. In general, however, no more than 5 lines of code are needed
6. Ignore all boxes that are completely empty; just focus on `# YOUR CODE HERE`

## Begin of the Test

#### Exercise 1: complete the following function, that is to print all integer numbers $m \in [0..10]$, using a for loop. The expected output is: 
```
integer number: 0
integer number: 1
integer number: 2
...
integer number: 10
```

11 lines in total, no tabs before, just lowercase, and a single space after the colon `:`

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
def print_numbers_in_for_loop(first, last):
    # for all integer numbers m in [first..last] ...
    # YOUR CODE HERE
    raise NotImplementedError()

print_numbers_in_for_loop(0,10)
```

Output:
```
<ipython-input-1-1392bffc3f33> in <module>
      4     raise NotImplementedError()
      5 
----> 6 print_numbers_in_for_loop(0,10)

<ipython-input-1-1392bffc3f33> in print_numbers_in_for_loop(first, last)
      2     # for all integer numbers m in [first..last] ...
      3     # YOUR CODE HERE
----> 4     raise NotImplementedError()
      5 
      6 print_numbers_in_for_loop(0,10)

NotImplementedError:
```

<!-- grade_id: at1, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import sys
import unittest
from contextlib import contextmanager
from io import StringIO

@contextmanager
def captured_output():
    new_out, new_err = StringIO(), StringIO()
    old_out, old_err = sys.stdout, sys.stderr
    try:
        sys.stdout, sys.stderr = new_out, new_err
        yield sys.stdout, sys.stderr
    finally:
        sys.stdout, sys.stderr = old_out, old_err

    
class MyEntryTestClasses(unittest.TestCase):

    def test_q1_print_numbers_for_loop(self):
        
        with captured_output() as (out, err):
            print_numbers_in_for_loop(0, 10)
    
        # This can go inside or outside the `with` block
        output = out.getvalue().strip()
        self.assertEqual(output, 
        'integer number: 0\n'
        'integer number: 1\n'
        'integer number: 2\n'
        'integer number: 3\n'
        'integer number: 4\n'
        'integer number: 5\n'
        'integer number: 6\n'
        'integer number: 7\n'
        'integer number: 8\n'
        'integer number: 9\n'
        'integer number: 10')
        
def run_MyEntryTestClasses():
    suite = unittest.TestLoader().loadTestsFromTestCase(MyEntryTestClasses)
    unittest.TextTestRunner(verbosity=2).run(suite)

run_MyEntryTestClasses()
### END HIDDEN TESTS
```

Output:
```
test_q1_print_numbers_for_loop (__main__.MyEntryTestClasses) ... ERROR

======================================================================
ERROR: test_q1_print_numbers_for_loop (__main__.MyEntryTestClasses)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "<ipython-input-2-555dc31e27be>", line 23, in test_q1_print_numbers_for_loop
    print_numbers_in_for_loop(0, 10)
  File "<ipython-input-1-1392bffc3f33>", line 4, in print_numbers_in_for_loop
    raise NotImplementedError()
NotImplementedError

----------------------------------------------------------------------
Ran 1 test in 0.001s

FAILED (errors=1)
```

#### Exercise 2: given a sorted list `my_list = [0,1,2,3,4,5,6,7,8,9,10]`. Also given: 4 initially empty lists `list_1, list_2, list_3, list_4` and 4 half-open intervals:

\begin{enumerate}
\item \hspace{1.0cm} $I_1 = [0..2)$

\item \hspace{1.0cm} $I_2 = [2..5)$   

\item \hspace{1.0cm} $I_3 = [5..8)$   

\item \hspace{1.0cm} $I_4 = [8..11)$   
\end{enumerate}

Copy each of the integer numbers in `my_list` belonging to interval $I_k$ to `list_k`, using `if-elif-else` statements. To copy a number `m` to `list_k` use: `list_k.append(m)`.

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
my_list = [0,1,2,3,4,5,6,7,8,9,10]
list_1, list_2, list_3, list_4 = [], [], [], []

# sub-optimal but most likely solution ...
for m in my_list:
    # put each m in the correct list ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
# print all lists...
for k, lst in enumerate([list_1, list_2, list_3, list_4]):
    print('list {:d}: {:s}'.format(k+1, str(lst)))
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-3-33a70802cea7> in <module>
      6     # put each m in the correct list ...
      7     # YOUR CODE HERE
----> 8     raise NotImplementedError()
      9 
     10 # print all lists...

NotImplementedError:
```

<!-- grade_id: at21, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_1),2)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-4-876999a56fc6> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(len(list_1),2)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at22, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_1), [0, 1])
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-5-6f94ea92fe0f> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(sorted(list_1), [0, 1])
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at23, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_2),3)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-6-872518fecd23> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(len(list_2),3)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at24, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_2), [2, 3, 4])
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-7-1c3ecb383d28> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(sorted(list_2), [2, 3, 4])
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at25, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_3),3)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-8-ec302f5f7809> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(len(list_3),3)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at26, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_3), [5, 6, 7])
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-9-5de4379d1d38> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(sorted(list_3), [5, 6, 7])
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at27, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(len(list_4),3)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-10-5357666b7955> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(len(list_4),3)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at28, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(sorted(list_4), [8, 9, 10])
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-11-cdd4c1f87c69> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(sorted(list_4), [8, 9, 10])
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

#### Exercise 3
Let $P > 0$ be an arbitrary big number. Recall that for a big enough $n$, we have that $P^{\left( \frac{1}{2} \right) ^n} \approx 1$. 

Repeatedly take the square root of $P$ until:

$|~P-1~| \leq \epsilon$

Here, $|P-1|$ denotes the absolute value of $P-1$. Use $\epsilon = 1.0E-06$. 

Complete the below skeleton.

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
import math

bigNum = 879857982742398758947958728934798758934

EPS = 1.0e-06
number = bigNum

n_root = 0
# while number more than EPS away from 1.0 do ...
# YOUR CODE HERE
raise NotImplementedError()
    number = math.sqrt(number)
    n_root += 1
    if n_root >= 50:
        break

    
print('the {:d}-th root of {:d} = 1.000000 +/- {:8.6e}'. format(n_root, bigNum, EPS))
```

Output:
```
File "<ipython-input-12-b9ee49600da3>", line 12
    number = math.sqrt(number)
    ^
IndentationError: unexpected indent
```

<!-- grade_id: at31, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# is the while loop broken?
assert n_root < 50
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-13-8476701714c6> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 # is the while loop broken?
----> 3 assert n_root < 50
      4 ### END HIDDEN TESTS

NameError: name 'n_root' is not defined
```

<!-- grade_id: at32, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(EPS,1.0e-06)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-14-16a7d327a518> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(EPS,1.0e-06)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at33, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_greater
assert_greater(bigNum, 0.0)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-15-b1c49ecf5e78> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_greater
      3 assert_greater(bigNum, 0.0)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at34, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_greater_equal
assert_greater_equal (EPS, abs(number-1.0))
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-16-75d616bbdb24> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_greater_equal
      3 assert_greater_equal (EPS, abs(number-1.0))
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at35, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# check whether big number and number are on the same side of 1
from nose.tools import assert_true
assert_true ( (bigNum > 1.0 and number > 1.0) or (bigNum < 1.0 and number < 1.0) )
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-17-16d7c836c687> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 # check whether big number and number are on the same side of 1
----> 3 from nose.tools import assert_true
      4 assert_true ( (bigNum > 1.0 and number > 1.0) or (bigNum < 1.0 and number < 1.0) )
      5 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

#### Exercise 4: let `my_list` be a list of items. Remove items from that list in a while loop, one-by-one, until you find that the list is empty. Print how many items the list apparently had. Do not use function `len()`; use  `item = my_list.pop()` to remove an item from the list, and increment the item count. Avoid popping from an empty list.

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
my_list = ['a', 1, 5, '6', (3, 5), 13, 11, 'Z', '!', 5, 2, 4, 9, 6, 'y']

item_count = 0
# while my_list is not empty do ...
# YOUR CODE HERE
raise NotImplementedError()
    item = my_list.pop()
    item_count += 1
    if item_count >= 50:
        break

print('my_list contained {:d} items'.format(item_count))
```

Output:
```
File "<ipython-input-18-76341394b7be>", line 7
    item = my_list.pop()
    ^
IndentationError: unexpected indent
```

<!-- grade_id: at41, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert len(my_list) == 0 # really empty?
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-19-ea1eb2d56851> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert len(my_list) == 0 # really empty?
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at42, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( item_count < 50 ) # loop not broken ?
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-20-387de18f40e1> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_true
      3 assert_true ( item_count < 50 ) # loop not broken ?
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at43, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(item_count,15)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-21-d82de6590678> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(item_count,15)
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

#### Exercise 5: Given a list `points` of 10 points $(x,y)$ in the plane. We select all points from `points` lieing in the first quadrant, i.e., above the $X$-axis ( line $y=0$ ) and right of the $Y$-axis ( line $x=0$ ) and collect them in a list `Q1`; the code skeleton below shows this. You are asked to select all other points (i.e., points lieing in quadrant two, three and four) by adding a second for-loop, with inverted  selection criteria, and collect these points in a single second list `Q234`. Print the list `Q234` as a string.

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
X0, Y0 = 0.0, 0.0

points = [(1,3), (-1,-2), (1,-1), (-2,0), (3,-1), (-4,1), (1,2), (-1,-3), (2,3), (-4,-1)]

Q1 = []
# make a for-loop to find points in quadrant 1 ...
for p in points:
    if p[0] > X0 and p[1] > Y0:
        Q1.append(p)
    
print('{:d} points in Q1  : {:s}'.format(len(Q1), str(Q1)))

Q234 = []
# now make a 2nd for-loop and invert the selection to find points in quadrant 2 to 4 ...
for p in points:
    # YOUR CODE HERE
    raise NotImplementedError()
    
# print the Q234 list ...
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
3 points in Q1  : [(1, 3), (1, 2), (2, 3)]
NotImplementedError                       Traceback (most recent call last)
<ipython-input-22-6a2c00344787> in <module>
     15 for p in points:
     16     # YOUR CODE HERE
---> 17     raise NotImplementedError()
     18 
     19 # print the Q234 list ...

NotImplementedError:
```

<!-- grade_id: at51, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
# check lieing in quadrant 1
for p in Q1:
    assert_true( p[0] > 0 and p[1] > 0 )
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-23-ad7154950c26> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_true
      3 # check lieing in quadrant 1
      4 for p in Q1:
      5     assert_true( p[0] > 0 and p[1] > 0 )

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at52, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( len(Q1) == 3 )
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-24-c1a036ccdfc5> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_true
      3 assert_true ( len(Q1) == 3 )
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at53, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
X0, Y0 = 0, 0
# check lieing in quadrant 2, 3, or 4
for p in Q234:
    assert_true( p[0] < 0 or p[1] < 0 )
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-25-f5ef86339ea0> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_true
      3 X0, Y0 = 0, 0
      4 # check lieing in quadrant 2, 3, or 4
      5 for p in Q234:

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at54, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true
assert_true ( len(Q234) == len(points) - len(Q1) )
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-26-57146940cfb0> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_true
      3 assert_true ( len(Q234) == len(points) - len(Q1) )
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

## End of the Test