# Python Entry Test

## Part 1 of 4: Variables (types, assignments, print formats, precision, operators)

## Instructions

Below are 5 Exercises: python problems you need to solve, in order to pass this test. For each correct answer, you will be assigned points, with a maximum of 100 points in total. If you score 80 points or higher, you have passed the test. You cannot see your score immediately; your test will be graded as part of a batch job, later on.

Note the following:
1. You are asked to code in **Python3** (version 3), not Python2
2. For some answers, you obtain a *skeleton*, to frame your answer. You are asked to complete it
3. The answer you are expected to given starts where it says:
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```
   Start your answer by *removing these two lines*. Do not put your answer anywhere else than *just here, in this box*
4. You may use as many lines for your answer as you see necessary
5. In general, however, no more than 5 lines of code are needed
6. Ignore all boxes that are completely empty; just focus on `# YOUR CODE HERE`
     

# Begin of the Test

#### Exercise 1: Function `pow(a,b)` computes `a` to the power `b`. Which of the following Python statements produces the result: 25 ? Set `option=`  below to the correct option

    1. print('{:d}'.format(1+pow(4,2)+8))
    2. print('{:d}'.format(4+pow(2,4)+8))
    3. print('{:d}'.format(1+pow(16)+8))
    4. print('{:d}'.format(1+pow(1,16)+8))

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
option=1
```

<!-- grade_id: at1, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert int(option) == 1
### END HIDDEN TESTS
```

#### Exercise 2: we have two variables: `a = 2.4` (a floating point number, type `float`) and `s = '0.01'` (a `string`). You are asked to create these variables as specified, using the code skeleton below. Next, you are asked to multiply `a` and ` s`, after turning `s` into a number, and assign their product to variable `product` (type `float`).

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
# correct the values of a and s ...
a = 2.4
s = '0.01'
# YOUR CODE HERE
product=a*float(s)
```

<!-- grade_id: at21, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(a, float)
### END HIDDEN TESTS
```

<!-- grade_id: at22, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(s, type('0.01'))
### END HIDDEN TESTS
```

<!-- grade_id: at23, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert product == a * float(s)
### END HIDDEN TESTS
```

<!-- grade_id: at24, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert a == 2.4
### END HIDDEN TESTS
```

<!-- grade_id: at25, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert s == '0.01'
### END HIDDEN TESTS
```

<!-- grade_id: at26, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert product == 0.024
### END HIDDEN TESTS
```

#### Exercise 3: create a list `my_list` of even numbers, starting at 2 up till and including 10. Next, compute the length of `my_list` using function `len()` and assign this length to variable `length`. 

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
my_list=[2,4,6,8,10]
length = len(my_list)
```

<!-- grade_id: at31, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(my_list, list)
### END HIDDEN TESTS
```

<!-- grade_id: at32, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(length,int)
### END HIDDEN TESTS
```

<!-- grade_id: at33, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert len(my_list) == 5
### END HIDDEN TESTS
```

<!-- grade_id: at34, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert sum(my_list) == 30
### END HIDDEN TESTS
```

#### Exercise 4: let `T = (1.0+0.1E-10, 1.0+1.0E-12)` be a 2-element tuple of numbers close to 1. Using the below skeleton, print the smallest and the greatest value (`min_value` and `max_value`, resp.) with 12 decimals of precision. You may use functions `min(), max()` to find the min and max values in `T`

<!-- grade_id: cell-q4, points: 0, grade: false, solution: true -->
```python
T = (1.0+0.1E-10, 1.0+1.0E-12)
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-15-5139b76e3bef> in <module>
      1 T = (1.0+0.1E-10, 1.0+1.0E-12)
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: at41, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
assert abs(min_value - T[1]) <= EPS_MIN
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-16-edbd5cc9175e> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
----> 3 assert abs(min_value - T[1]) <= EPS_MIN
      4 ### END HIDDEN TESTS

NameError: name 'min_value' is not defined
```

<!-- grade_id: at42, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
assert abs(max_value - T[0]) <= EPS_MAX
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-17-7e77d0d35b0a> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
----> 3 assert abs(max_value - T[0]) <= EPS_MAX
      4 ### END HIDDEN TESTS

NameError: name 'max_value' is not defined
```

<!-- grade_id: at43, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
assert round(min_value,12) == 1.000000000001
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-18-c5b58816e439> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
----> 3 assert round(min_value,12) == 1.000000000001
      4 ### END HIDDEN TESTS

NameError: name 'min_value' is not defined
```

<!-- grade_id: at44, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
assert round(max_value,12) == 1.000000000010
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-19-93faaa85efb8> in <module>
      1 ### BEGIN HIDDEN TESTS
      2 EPS_MIN, EPS_MAX = 1.0E-12, 1.0E-11
----> 3 assert round(max_value,12) == 1.000000000010
      4 ### END HIDDEN TESTS

NameError: name 'max_value' is not defined
```

#### Exercise 5: Given two lines of text: 

 1. 'This is our header line'
 2. 'Our paragraph of text starts here.'
 
Print the two lines with a single blank line between them, using the below skeleton. To do so, compose a single string `my_text` containing these three lines, so that they can be printed with a single `print()` statement. Make all characters **in the header line** uppercase (use `header.upper()` to do that). 

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
header = 'This is our header line'
pgraph =  'Our paragraph of text starts here.'
# YOUR CODE HERE
my_text = header.upper()+"\n"+pgraph
print('{:s}'.format(my_text))
```

Output:
```
THIS IS OUR HEADER LINE
Our paragraph of text starts here.
```

<!-- grade_id: at51, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(header, 'This is our header line')
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-21-15a91ca0ea4e> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(header, 'This is our header line')
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at52, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(pgraph, 'Our paragraph of text starts here.')
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-22-f450c62e422a> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 assert_equal(pgraph, 'Our paragraph of text starts here.')
      4 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at53, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
expected = header.upper() + '\n\n' + pgraph
assert(my_text.startswith('THIS IS OUR HEADER'))
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-23-4aabe2746dc7> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 expected = header.upper() + '\n\n' + pgraph
      4 assert(my_text.startswith('THIS IS OUR HEADER'))
      5 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: at54, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
expected = header.upper() + '\n\n' + pgraph
assert_equal(my_text, expected)
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-24-a9e86522fcfb> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 from nose.tools import assert_equal
      3 expected = header.upper() + '\n\n' + pgraph
      4 assert_equal(my_text, expected)
      5 ### END HIDDEN TESTS

ModuleNotFoundError: No module named 'nose'
```

# End of the test